<a href="https://colab.research.google.com/github/sikoh/DS-Linear-Algebra/blob/main/Inference-Linear-Regression/DS_InferenceForRegression_Homework.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Unit 1 Sprint 3 Module 2

## Module Project: Inference for Linear Regression

### Learning Objectives

* identify the appropriate hypotheses to test for a statistically significant relationship between two quantitative variables
* conduct and interpret a t-test for the slope parameter
* identify the appropriate parts of the output of a linear regression model and use them to build a confidence interval for the slope term.
* make the connection between the t-test for a population mean and a t-test for a slope coefficient.
* identify violations of the assumptions for linear regression

### Total notebook points: 16

## Introduction

### Statistical significance between head size and brain weight in healthy adult humans

The `Brainhead.csv` dataset provides information on 237 individuals who were subject to post-mortem examination at the Middlesex Hospital in London around the turn of the 20th century. Study authors used cadavers to see if a relationship between brain weight and other more easily measured physiological characterizes such as age, sex, and head size could be determined.

The end goal was to develop a way to estimate a person’s brain size while they were still alive (as the living aren’t keen on having their brains taken out and weighed).

**We wish to determine if there is a linear relationship between head size and brain weight in healthy human adults.**

Source: R.J. Gladstone (1905). "A Study of the Relations of the Brain to the Size of the Head", Biometrika, Vol. 4, pp105-123.


**Use the above information to complete Tasks 1 - 10.**



Recall from the Module 1 Project that we were working with the brain weight variable (`Brain`) and the head size (`Head`) variables. We identified the dependent and independent variables, plotted our variables on the appropriate axes, and then described the strength of the relationship.

* `Brain` (brain weight in g) - **dependent** variable (y)
* `Head` (the head size in cubic cm - **independent** variable (x)

Now, we're going to bring back some statistics from Sprint 2 and look at the statistical association between head size and brain weight.



**Task 1** - Statistical tests

What type of statistical test will we use to determine if there is a statistically significant association between head size and brain weight in the population?  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: A beta test

B: A chi-square test

C: A Bayesian test

D: A t-test


<!-- grade_id: cell-07bbb2a930fdc4a8, points: 0, grade: false, solution: true -->
```python
# Task 1

# YOUR CODE HERE
#raise NotImplementedError()
Answer = 'D'
```

**Task 1 Test**

<!-- grade_id: cell-926f6d1e0cfe693d, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 2** - Hypotheses for linear regression

Select the null and alternative hypotheses you would use to test for a statistically significant association between head size and brain weight.  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: $H_0: \beta_1 = 0$ vs. $H_a: \beta_1 \neq 0$

B: A$H_0: \beta_0 = 0$ vs. $H_a: \beta_0 \neq 0$

C: $H_0: \beta_1 \neq 0$ vs. $H_a: \beta_1 = 0$

D: $H_0: \beta_0 = \beta_1 $ vs. $H_a: \beta_0 \neq \beta_1$


<!-- grade_id: cell-57418db46031ec24, points: 0, grade: false, solution: true -->
```python
# Task 2

# YOUR CODE HERE
#raise NotImplementedError()
Answer = 'A'
```

**Task 2 Test**

<!-- grade_id: cell-4d7da361b388c8fa, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 3** - Load the data

As we usually begin, let's load the data! The URL has been provided.

* load your CSV file into a DataFrame named `df`

<!-- grade_id: cell-dce60d72fa032dbb, points: 0, grade: false, solution: true -->
```python
# Task 3

# Imports
import pandas as pd
import numpy as np

data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Brainhead/Brainhead.csv'

# YOUR CODE HERE
#raise NotImplementedError()
df = pd.read_csv(data_url)
# Print out your DataFrame
df.head()
```

Output:
```
Gender  Age  Head  Brain
0       1    1  4512   1530
1       1    1  3738   1297
2       1    1  4261   1335
3       1    1  3777   1282
4       1    1  4177   1590
```

```python
df.shape
```

Output:
```
(237, 4)
```

**Task 3 Test**

<!-- grade_id: cell-d4edaaeebedaf8ff, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named `df`?'
assert len(df) == 237

```

**Task 4** - Fit OLS model

Now, we're going to fit a regression model to our two variables. We're going to use `statsmodels.formula.api` and import the `ols` model. This import has been provided for you.

* Fit a model and name your variable `model`
* Using the `model.params[1]` method, assign the slope to the variable `slope`. Your variable should be a float (`numpy.float64`).
* Using the same `model.params[0]` method, assign the intercept to the variable `intercept`. Your variable should be a float (`numpy.float64`).
* Print out your model summary.

*Hint: Make sure to use the format Y ~ X for the model input.*

<!-- grade_id: cell-e9560546a99c90ac, points: 0, grade: false, solution: true -->
```python
# Task 4

# Import statsmodels - DON'T Delete
from statsmodels.formula.api import ols

# Set-up and fit the model in one step
# (format Y ~ X)
model = ols('Brain ~ Head', data = df).fit()
# YOUR CODE HERE
#raise NotImplementedError()
intercept, slope = model.params
print (intercept, slope)
# Print the model summary
print(model.summary())
```

Output:
```
325.573421049442 0.2634293394893992
                            OLS Regression Results                            
==============================================================================
Dep. Variable:                  Brain   R-squared:                       0.639
Model:                            OLS   Adj. R-squared:                  0.638
Method:                 Least Squares   F-statistic:                     416.5
Date:                Tue, 24 Oct 2023   Prob (F-statistic):           5.96e-54
Time:                        13:59:03   Log-Likelihood:                -1350.3
No. Observations:                 237   AIC:                             2705.
Df Residuals:                     235   BIC:                             2711.
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                
```

**Task 4 - Test**

<!-- grade_id: cell-100a97e8d1184783, points: 2, grade: true, solution: false -->
```python
# Task 4 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - Formulate the statistical model

Using the model parameters returned above, you will now write out the statistical model as a linear equation. Remember, we are predicting brain weight from head size.

* write your equation in Python
    * assign the dependent variable to `y_hat`
    * assign the independent variable to `x` with a value of `4000`
    * use the variables you created earlier for the slope and intercept

<!-- grade_id: cell-1ddf887e0e3e603e, points: 0, grade: false, solution: true -->
```python
# Task 5
x = 400
# YOUR CODE HERE
y_hat = intercept + slope*x
#raise NotImplementedError()
print(y_hat)
```

Output:
```
430.94515684520167
```

**Task 5 - Test**

<!-- grade_id: cell-f73a4aaf70997ef3, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6** - Statistical parameters

Now that we have fit a model, we're going to pull out the statistical parameters.

* assign the standard error to the variable `std_err`
* assign the value of the t-statistics to the variable `t_stat`
* assign the p-value for the slope to the variable `p_slope`

**Assign values out to the 1/1000 place (for example, `777.555`)**

<!-- grade_id: cell-0088cec3b26fca73, points: 0, grade: false, solution: true -->
```python
# Task 6

# YOUR CODE HERE
#raise NotImplementedError()
std_err = model.bse[1]
print(std_err)
t_stat = model.tvalues[1]
print(t_stat)
p_slope = model.pvalues[1]
print(p_slope)
```

Output:
```
0.012907433440886976
20.40911856689743
5.957630839406541e-54
```

**Task 6 - Test**

<!-- grade_id: cell-3b6f76045e89854a, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 7** - Hypothesis Test

Conduct your hypothesis test and select the answer below that best describes the conclusion at the 0.05 significance level.  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: We reject the null hypothesis and conclude there is no association between head size and brain weight.

B: We reject the null hypothesis and conclude there is an association between head size and brain weight.

C: We fail to reject the null hypothesis and conclude there is no association between head size and brain weight.

D: We fail to reject the null hypothesis and conclude there is an association between head size and brain weight.


<!-- grade_id: cell-8a9fb0cf3d1b46d7, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
#raise NotImplementedError()
Answer = 'B'
```

**Task 7 Test**

<!-- grade_id: cell-3970fbb1a3eaefb5, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Hypothesis Test for the intercept

Conduct a hypothesis test for the intercept and select the answer below that best describes the conclusion at the 0.05 significance level.  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: We reject the null hypothesis and conclude there is no association between head size and brain weight.

B: We fail to reject the null hypothesis and conclude there is no association between head size and brain weight.

C: It does not make sense to interpret or conduct a hypothesis test for the intercept because 0 is not a value of x in our dataset.

D: It does not make sense to interpret or conduct a hypothesis test for the intercept because the slope is statistically significant.


<!-- grade_id: cell-a8ea8470352ac330, points: 0, grade: false, solution: true -->
```python
# Task 8

# YOUR CODE HERE
raise NotImplementedError()

```

**Task 8 Test**

<!-- grade_id: cell-a3a3433557e73bac, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9** - Confidence Interval

Calculate the 95% confidence interval for your slope term. Use your model summary to find these values. Assign them values out to the 1/1000 place (for example, `ci_low = 0.345`)

* assign the lower value of brain size to `brain_low`
* assign the upper value of brain size to `brain_upper`


<!-- grade_id: cell-cebcd0240ee16185, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 9 - Test**

<!-- grade_id: cell-9a0563abb4c16599, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Plot confidence intervals

Plot the confidence limits on both the slope and intercept terms with a shaded area around our regression line. The easiest way to do this is using `seaborn` with the `ci` parameter set to the confidence level you want (the default is `ci=95`). This plot will not be autograded.

<!-- grade_id: cell-be916b022e3948d5, points: 0, grade: false, solution: true -->
```python
# Task 10

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 10** - Correlation (short answer)

Does it seem plausible that larger head size causes greater brain weight?  Or is it possible that something else causes differences in both of those factors?

ANSWER:



## Part B

### Sleep Data

Use the following information to answer Tasks 11 - 18 in the rest of this project:

Researchers recorded data on sleep duration as well as a set of ecological and constitutional variables for a selection of mammal species. This data is available in the Sleep.csv dataset; the URL is provided below.

(*Source: Allison, T. and Cicchetti, D. (1976), "Sleep in Mammals: Ecological and Constitutional Correlates",  Science, November 12, vol. 194, pp. 732-734.*)

**Data Dictionary:**

| Variable Name |            Description           |                                                                 Details                                                                 |              |   |
|:-------------:|:--------------------------------:|:---------------------------------------------------------------------------------------------------------------------------------------:|--------------|---|
| Animal        |            Animal name           |                                                            Character [string]                                                           |    |   |
| Body          |            Body weight           |                                                               Numeric [kg]                                                              |       |   |
| Brain         |           Brain weight           |                                                               Numeric [g]                                                               |  |   |
| SWsleep       | Slow wave (“non-dreaming”) sleep | Numeric [hours]                                                                                                                         |              |   |
| Parasleep     | Paradoxical (“dreaming”) sleep   | Numeric [hours]                                                                                                                         |              |   |
| Totsleep      | Total sleep                      | Numeric [hours]                                                                                                                         |              |   |
| Life          | Maximum life expectancy          | Numeric [years]                                                                                                                         |              |   |
| Gest          | Gestation time                   | Numeric [days]                                                                                                                          |              |   |
| Pred          | Predation index                  | Numeric [1 – 5] 1 = least likely to be preyed upon, 5 = most likely to be preyed upon                                                   |              |   |
| Exposure      | Sleep exposure index             | Numeric [1 – 5] 1 = least amount of exposure during sleep (mammal sleeps indoors or in a den), 5 = most amount of exposure during sleep |              |   |
| Danger        | Overall danger index             | Numeric [ 1 – 5] 1 = least amount of danger from other animals, 5 = most amount of danger from other animals                            |              |   |



**Task 11**

Before we can look at the data, we need to load in the data. The URL has been provided.

* load in the CSV file as a DataFrame and assign it to the variable `df_sleep`
* make sure to view the DataFrame!

<!-- grade_id: cell-8b27985340fda74e, points: 0, grade: false, solution: true -->
```python
# Task 11

data_url_2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Sleep/Sleep.csv'

# YOUR CODE HERE
raise NotImplementedError()

# Print out your DataFrame
df_sleep.head()
```

**Task 11 - Test**

<!-- grade_id: cell-728b04417127d100, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

assert isinstance(df_sleep, pd.DataFrame), 'Have you created a DataFrame named `df_sleep`?'
assert len(df_sleep) == 42

```

**Task 12** - Plot to check for linearity

Plot the relationship between *gestation time* and time spent in *dreaming sleep*. This plot will not be autograded.

* you can use `seaborn` for your plot, with the `regplot()` or `lmplot()`
* include the regression line but turn off the confidence interval (`ci=None`)

Select the answer that best describes the relationship between gestation time and time spent in dreaming sleep. Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: Gestation time and time spent in dreaming sleep appear to have a linear relationship.  The linear regression line shown in the figure seems to fit the data well.

B: Gestation time and time spent in dreaming sleep appear to have a non-linear relationship.  The linear regression line shown in the figure does not seem to fit the data well.

C: Gestation time and time spent in dreaming sleep appear to have no relationship.  The linear regression line shown in the figure seems to fit the data well.

D: Gestation time and time spent in dreaming sleep appear to have a linear relationship.  The linear regression line shown in the figure does not seem to fit the data well.

<!-- grade_id: cell-6b9a56b997acc63d, points: 0, grade: false, solution: true -->
```python
# Task 12

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 12 Test**

<!-- grade_id: cell-9ac8291f88c21441, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 13** - Transform a variable

Let's try something new: taking the log of a variable to transform it. Then we'll look at the relationship between the log of that variable and the other variable (which will remain the same)

* Create a new variable with the log of gestational time and add it as a column to `df_sleep` with the name `log_gest`

*Hint: use the natural log `np.log()`*

<!-- grade_id: cell-d811349ababa8892, points: 0, grade: false, solution: true -->
```python
# Task 13

# YOUR CODE HERE
raise NotImplementedError()

# Look at your new column
df_sleep.head()
```

**Task 13 - Test**

<!-- grade_id: cell-94ac629ec8bbe280, points: 1, grade: true, solution: false -->
```python
# Task 13 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 14** - Plot a new relationship

Plot the relationship between *log of gestation time* and time spent in *dreaming sleep*. This plot will not be autograded.

* you can use `seaborn` for your plot, with the `regplot()` or `lmplot()`
* include the regression line but turn off the confidence interval (`ci=None`)

Select the answer that best describes the relationship between the log of gestation time and time spent in dreaming sleep. Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: The log of gestation time and time spent in dreaming sleep appear to have a linear relationship.  The linear regression line shown in the figure seems to fit the data well.

B: The log of gestation time and time spent in dreaming sleep appear to have a non-linear relationship.  The linear regression line shown in the figure does not seem to fit the data well.

C: The log of gestation time and time spent in dreaming sleep appear to have no relationship.  The linear regression line shown in the figure seems to fit the data well.

D: The log of gestation time and time spent in dreaming sleep appear to have a linear relationship.  The linear regression line shown in the figure does not seem to fit the data well.

<!-- grade_id: cell-6607d864b2e34089, points: 0, grade: false, solution: true -->
```python
# Task 14

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 14 Test**

<!-- grade_id: cell-4d42d4ebe8c5c6d1, points: 1, grade: true, solution: false -->
```python
# Task 14 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 15** - Model the sleep data

Next, create a model of the relationship between the log of gestation time and dreaming sleep using the `statsmodels.formula.api`. Remember that the `statsmodels` import was made earlier.

* Fit a OLS model and assign it to the variable `model_sleep` (remember to enter the model in the format Y ~ X).
* Print out your model summary.
* Answer the questions below to interpret your results.

<!-- grade_id: cell-a08ff1b8439746f9, points: 0, grade: false, solution: true -->
```python
# Task 15

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary
print(model_sleep.summary())
```

**Task 15 Test**

<!-- grade_id: cell-dbf3e67eb3732036, points: 1, grade: true, solution: false -->
```python
# Task 15 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 16** - Statistical significance of sleep data (short answer)

Is the *log of gestational time* statistically significantly associated with time spent in *dreaming sleep* at the alpha = 0.05 level?

ANSWER:

**Task 17** - Predicting dreaming sleep from gestation time

The final task! Using the model we just created, predict the amount of dreaming sleep for a mammal that gestates her young for 262 days (this is the gestation time).

* Assign the gestation time to the variable `x_predict`. This variable will be an integer.
* Take the log of `x_predict` and assign it to the variable `ln_x_predict`. THis variable will be a float.
* Use the `slope` and `intercept` variables from your `model_sleep` to complete the calculation.
* Your result should be a float and assigned to the variable `sleep_predict`.

<!-- grade_id: cell-b516863ad6936f72, points: 0, grade: false, solution: true -->
```python
# Task 17

# YOUR CODE HERE
raise NotImplementedError()

# Print out the log of x and the predicted sleep value
print('ln 262 = ', ln_x_predict)
print('Predicted dreaming sleep = ', sleep_predict)
```

**Task 17 Test**

<!-- grade_id: cell-772974b44d216ef9, points: 1, grade: true, solution: false -->
```python
# Task 16

assert ln_x_predict == np.log(x_predict), 'Did you use the correct log calculation?'


```